# Artificial Intelligence for Robotics 04

## General Information:
Please do not add or delete any cells. Answers belong into the corresponding cells (below the question). If a function is given (either as a signature or a full function), you should not change the name, arguments or return value of the function.

If you encounter empty cells underneath the answer that can not be edited, please ignore them, they are for testing purposes.

When editing an assignment there can be the case that there are variables in the kernel. To make sure your assignment works, please restart the kernel and run all cells before submitting (e.g. via Kernel -> Restart & Run All).

Code cells where you are supposed to give your answer often include the line `raise NotImplementedError`. This makes it easier to automatically grade answers. If you edit the cell please outcomment or delete this line.

## Submission:
Please submit your notebook via the web interface (in the main view -> Assignments -> Submit). The assignments are **due on Monday at 20:00**.

## Group Work:
Please enter your UID (your username here) and those of your group partners into the next cell. We apply plagiarism checking, so do not submit others solutions! If an assignment has a copied solution, the task will be graded with 0 points for all people with the same solution.

## Questions about the Assignment:
If you have questions about the assignment please post them in the LEA forum before the deadline. Don't wait until the last day to post questions!

### Please add the usernames of all your team members in the manner member1, member2 in next cell (example given below)

member1 = 'example'

member2 = 'example2'

#### If you are not working in a group, then please add member2 as none2s

<!-- grade_id: cell-a9c2f4587c671f98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e27ced94e387fe18, points: 0, grade: true, solution: false -->
```python
# Execute this cell to make sure you correctly filled in the usernames of the team members

def group_name_test():
    for member_id in [member1, member2]:
        assert isinstance(member_id, str), "Please give your member id as a string."
        assert len(member_id) > 0, "You need to fill in the member id for both members"
        assert member_id.endswith("2s"), "The member id should end with 2s (Your JupyterHub username)"

group_name_test() 
print("All tests passed!")
```

# Task 1

**[18 Point(s)]**


# A* Theory [20 Points]

During the lecture, you have learned about the A* search algorithm and its propoerties. Please answere the following questions short and consice and be as precise as possible.


## Task 1.1

**[2 Point(s)]**

### Completeness [2 Points]

How is the *completeness* of an algorithm defined? Is A* complete?

<!-- grade_id: Astar_theory_A, points: 2, grade: true, solution: true -->
Completeness of algorithm is the ability to find a solution for every valid instance of the problem regardless whether it's the optimal solution or not, within a finite amount of time.

A* is a complete algorithm under certain conditions : The graph being search is finite, The hurestic used is admissible.

## Task 1.2

**[2 Point(s)]**

### Soundness [2 Points]

When is an algorithm be considered *sound*? Is A* sound?

<!-- grade_id: Astar_theory_B, points: 2, grade: true, solution: true -->
An algorithm is considered sound when it always produces the correct result for every valid input. A sound algorithm is the one that guarantees correctness.

A* is a sound algorithm if it is implemented correctly, it must satisfy two criterias: admissibility of Heuristic, correctness of implementation.

## Task 1.3

**[2 Point(s)]**

### Termination [2 Points]

When does A* end the search process?

<!-- grade_id: Astar_theory_C, points: 2, grade: true, solution: true -->
A* ends the search process either by finding a goal node or exhausting all possible paths to explore without finding a solution.

## Task 1.4

**[5 Point(s)]**

### Optimality [5 Points]

What does it mean when we say an algorithm is *optimal*? What are the criteria that have to be met in order for A* to be *optimal*? Provide an argument or proof on why this is required.

<!-- grade_id: Astar_theory_D, points: 5, grade: true, solution: true -->
An algorithm is optimal when it finds the best solution according to a specific criterion among all possible solutions. The criterion could be minimizing or maximizing a certain value (cost, time, distance). 

A* can be optimal under certain conditions, particularly when it satisfies the following criteria:
- Admissible Heuristic : heuristic guides A* to explore the most promising path first, ensuring that it reaches the goal along the shortest path. If the heuristic function is admissible it never overestimates the actual cost to reach the goal from any given node.
- Consistent Heuristic : if the heuristic function is also consistent then A* is not only optimal but also more efficient, a consistent heuristic ensures that A* expands nodes in the optimal order and avoids re-evaluating nodes unnecessarily, leading to improved efficiency.

## Task 1.5

**[5 Point(s)]**

### Time Complexity [5 Points]

What is the time complexity of A*? Explain, or derive mathematically! Just writing $O(...)$ is not sufficient!

<!-- grade_id: Astar_theory_E, points: 5, grade: true, solution: true -->
time complexity of A* can vary depending on several factors (size of search space, the quality of the heuristic function, and the implementation details), in the worst-case scenario, the time complexity of A* is exponential.

In the worst case, where the search space is very large and the heuristic function is not effective at guiding the search, A* may need to explore a large portion of the search space before finding the goal. In this case, the time complexity can be exponential in terms of both the number of nodes expanded and the depth of the search tree.

Explaination : 
In the worst case, the time complexity of A* can be analyzed by considering the number of nodes expanded during the search.
- b: The average branching factor of the search tree (i.e., the average number of successor nodes for each node).
- d: The depth of the optimal solution.
- h*(n): The actual cost from node n to the goal node.
- h(n): The heuristic estimate of the cost from node n to the goal node.
A*(n) expands nodes in the order of increasing f(n)=g(n)+h(n), where:
    g(n) is the cost from start node to node n.
    h(n) the heuristic estimate of the cost from node n to the goal node.

    - in the worst case, A* may need to expand all node up to depth (d) to find the optimal solution. Thus,the worst case time complexity can be approximated by "considering the number of nodes expanded up to depth d" 
    - if the heurstic function is admissible, then h(n) <= h*(n) for all nodes n. Therefore, A* expanding nodes in non-decreasing order of the actual cost g(n)+h*(n)
    - let's assume C* as the cost of the optimal solution, since h(n)  is an underestimate of the true cost, g(n)+h*(n) is an overestimate of C* for any node n.
    - In the worst case, A* expands nodes until it reaches the goal node, which has a cost C* . since h(n) is an underestimate, A* may explore nodes with f(n) > C* but it doesn't expand nodes with f(n) > C* because C* is is the minimum possible cost to reach the goal. Therefore, the number of nodes expanded by A* in the worst case is bounded by the number of nodes with f(n) <= C* up to depth d.
    - The numbers of nodes at depth d in a tree with an average branching factor of b is O(b^d), the worst case time complexity of A* is O(b^d).

## Task 1.6

**[2 Point(s)]**

### Consistency of a Heuristic [2 Points]

How is a consistent heuristic defined? Why is it a critical property in the context of the A* search algorithm?

<!-- grade_id: Astar_theory_F, points: 2, grade: true, solution: true -->
consistent heuristic is when the algorithm satisfes the condition of consistency which is defined as follow: A heuristic function h(n) is consistent if, for every node n and every successor n' of n generated by any action a, he estimated cost of reaching the goal from n is no greater than the cost of getting to n' plus  the estimated cost of reaching the goal from n' under the heuristic function.

Consistnecy is critical for A* because it ensures the optimality of the solution found by the algorithm.

# Task 2

**[50 Point(s)]**

# Applying Informed Search [60 Points]

Informed methods help us gain information about solving a problem through its current state space. This keeps a program from blundering around blindly guessing. Informed search strategies make use of information about the path of moves we took to get where we are currently by using an *evaluation function*. This evaluation function often makes use of a *heuristic function*, that provides an estimate of the distance of a node to the goal in a relaxed version of the problem.

# A Real World Problem

Consider a robot that works in a warehouse. The robot can carry packages on top of itself on a set of omnidirectional conveyors. It can carry a maximum of 8 packages in a 3x3 grid, with one spot left open. There is a slot from which the robot can dispense packages onto shelves. The robot must deposit the packages in order, so they must be lined up in a certain order near the slot. However, workers place the packages onto the robot randomly. Therefore, the robot must first sort the packages before depositing them. It can use the free space and the conveyors to shift the packages horizontally and vertically. 

## This becomes what is known as the 8 puzzle.

## Task 2.1

**[15 Point(s)]**

## The 8 Puzzle [10 Points]

An 8 puzzle is a simple game consisting of a 3 x 3 grid (containing 9 squares). One of the squares is empty.
The object is to move the squares around into different positions and having the numbers displayed in the "goal state". The squares can obviously only move along the horizontal and vertical axis and only into the **empty tile**. Therefore the problem can also be viewed as the empty tile "trading" positions with one of its neighbouring tiles.

The image below depicts the goal state of an "3 x 3" 8 puzzle.
<img src="Informed_Search_files/img/8_puzzle.png" alt="8 Puzzle" width="200"/>

The complexity of possible moves toward the final solution in a game like this is large. With an uninformed search strategy, finding a solution to this problem can take a lot of time. The problem is now about finding the shortest path to the goal state within reasonable time, which we will do in this assignment using A* and Greedy Search.

### Task
- We will implement the 8 puzzle game as the class `Puzzle`
- Define methods for moving a tile around in the puzzle
- Define a mehod for generating all possible moves from one particular configuration.
- You are already provided with a class and function templates below. Please complete the implementation.

<!-- grade_id: Astar_implementation_A, points: 0, grade: false, solution: true -->
```python
from typing import List, Tuple
from utils import print_puzzle

class Puzzle:

    def __init__(self, init_state: List[int], puzzle_type: int = 8):
        """
        Initialize a new 8-puzzle board
        :param init_state: Initial configuration of the board
        """
        self.init_state = init_state
        self.goal_state = [i for i in range(1, 9)] + [0]
        # YOUR CODE HERE
        raise NotImplementedError()

    def goal_test(self, state: List[int]):
        """Test if goal state is reached
        :param state: board configuration to check
        :return: true if the passed configuration is equal to goal configuration
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def is_explored(self, state: List[int]):
        """Check if a particular board configuration has already been explored
        :param state: board configuration to check
        :return: true if a particular configuration has already been explored
        """
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def reset(self):
        # YOUR CODE HERE
        raise NotImplementedError()


def move_left(position: int) -> int:
    """Move one position left in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving to the left
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def move_right(position: int) -> int:
    """Move one position right in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving to the right
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def move_up(position: int) -> int:
    """Move one position up in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving upwards
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def move_down(position: int):
    """Move one position down in 8 puzzle if possible
    :param position: current position of the 0 tile.
    :return: new position of the 0 tile after moving downwards
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def get_possible_moves(state: List[int]) -> List[List[int]]:
    """Check whether a move is possible in left, right, up, down direction and store it.
    :param state: current configuration of the puzzle as one dimensional list
    :return: list containing the new configurations after applying all possible moves
    """
    # YOUR CODE HERE
    raise NotImplementedError()

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_Astar_implementation_A0, points: 15, grade: true, solution: false -->
```python

```

## Task 2.2

**[5 Point(s)]**

## Heuristics [5 Points]

In this part, you are tasked to implement two different admissible heuristics for the 8 puzzle: The [Hamming distance](https://en.wikipedia.org/wiki/Hamming_distance) and the [Manhattan Distance](https://en.wikipedia.org/wiki/Taxicab_geometry).

**Hamming Distance**  
Let $A$ be an alphabet of symbols and $C$ a subset of $A^n$, the set of worlds of length $n$ over $A$. Let $u = (u_1, ..., u_n)$ and $v = (v_1, ...,v_n)$. The Hamming distance is deinfed as the number of places in which $u$ and $v$ differ. That is:
$$h(n) = \{i:u_i \neq v_i,i=1,...,n\}$$
See also: [Encyclopedia of Math](https://encyclopediaofmath.org/wiki/Hamming_distance)

It is clear that this heuristic is admissible since the total number of moves to order the tiles correctly is at least the number of misplaced tiles (each tile not in place must be moved at least once). The cost (number of moves) to the goal (an ordered puzzle) is at least the Hamming distance of the puzzle.

**Manhattan Distance**  
The Manhattan distance between two points $x = (x_1, ..., x_n)$ and $y = (y_1, ..., y_n)$ in $n$-dimensional space is the sum of the distances in each dimension.
$$d(x,y) = \sum^n_{i=1} \vert x_i - y_i \vert$$
See also: [Encyclopedia of Machine Learning](https://link.springer.com/referenceworkentry/10.1007%2F978-0-387-30164-8_506)

The Manhattan distance is an admissible heuristic because every tile will have to be moved at least the number of spots in between itself and its correct position.

### Task
- You are provided with the function definitions for the two heuristics below
- Please implement them according to the mathematical definition given above

<!-- grade_id: Astar_implementation_B, points: 0, grade: false, solution: true -->
```python
from typing import Tuple, List

def hamming_distance(state: List[int]) -> int:
    """Calculate the Hamming distance for a particular configuration
    :param state: current configuration of the puzzle
    :return: the number of misplaced tiles in the given configuration
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def manhattan_distance(state: List[int]) -> int:
    """Function to calculate the manhattan distance for a
    particular configuration
    :param state: current configuration of the puzzle
    :return: the accumulated manhattan distance between each tile and its goal position in the given configuration
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Run this test cell in order to get an indication of whether your implementation is working as expected.
assert hamming_distance([0,1,2,3,4,5,8,6,7]) == 9
assert manhattan_distance([1,2,3,4,0,5,6,8,7]) == 8
```

<!-- grade_id: test_Astar_implementation_B0, points: 5, grade: true, solution: false -->
```python

```

## Task 2.3

**[10 Point(s)]**

## A* Search Algorithm [10 Points]

Now you will implement the A* search algorithm. To do that, you can consider the pseudocode from the lecture

```
function A*(problem) returns a solution, or failure
    closed ← an empty set
    fringe ← a list containing Make-Node(Initial-State[problem])
    loop do
        if fringe is empty then return failure
        node ← Remove-Front(fringe)
        if Goal-Test[problem] applied to State(node) succeeds return node
            insert node into closed
        for each node n ∈ Expand(node, problem) do
            if there is a node m ∈ closed ∪ fringe such that
                STATE(m) = STATE(n) and f(m) ≤ f(n)
            then do nothing
            else
                insert n into fringe after the last node m such that f(m) ≤ f(n)
end
```

### Task
- You are provided with the function definition for the A* function
- Please implement A* according to the algorithm design discussed in the lectures

In order to get an estimation on whether your solution work as expected, there is a public test available in the test cell. Make sure that all dependent cells are executed first, otherwise some functions or variables might not exist and an error will occur.

<!-- grade_id: Astar_implementation_C, points: 0, grade: false, solution: true -->
```python
from typing import Callable, List, Tuple

def astar_search(board: Puzzle, heuristic: Callable) -> Tuple[List[int], int]:
    """
    :param board: the 8-puzzle to solve
    :param heuristic: the heuristic function to use
    :return: an ordered list with the solution path and the number of total nodes expanded
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Run this test cell in order to get an INDICATION of whether your implementation is working as expected.
import time
t0 = time.time()
p = Puzzle([0,1,2,3,4,5,8,6,7])
path, expanded_nodes = astar_search(p, manhattan_distance)
t1 = time.time()
assert len(path) == 23
# The execution time is limited to 10min, so please keep this in mind when implementing your solution!
assert t1-t0 <= 600
```

<!-- grade_id: test_Astar_implementation_C0, points: 10, grade: true, solution: false -->
```python

```

## Task 2.4

**[10 Point(s)]**

## Greedy Search [10 Points]

Now we will implement a *greedy search algorithm*. Recall from the lecture, that the evaluation function $f(n)$ is defined as $f(n) = h(n)$, with $h(n)$ being the heuristic function.

### Task
- You are provided with the definition of the greedy search function
- Implement a greedy search strategy by completing the function below

In order to get an estimation on whether your solution work as expected, there is a public test available in the test cell. Make sure that all dependent cells are executed first, otherwise some functions or variables might not exist and an error will occur.

<!-- grade_id: Astar_implementation_D, points: 0, grade: false, solution: true -->
```python
from typing import List, Tuple, Callable

def greedy_search(board: Puzzle, heuristic: Callable) -> Tuple[List[int], int]:
    """Implementation of the greedy search algorithm.
    :param board: the 8-puzzle to solve
    :param heuristic: the heuristic function to use
    :return: an ordered list with the solution path and the number of total nodes expanded 
    """
    # YOUR CODE HERE
    raise NotImplementedError()

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Run this test cell in order to get an INDICATION of whether your implementation is working as expected.
import time
t0 = time.time()
p = Puzzle([0,1,2,3,4,5,8,6,7])
path, expanded_nodes = greedy_search(p, manhattan_distance)
t1 = time.time()
# The execution time is limited to 10min, so please keep this in mind when implementing your solution!
assert t1-t0 <= 600
```

<!-- grade_id: test_Astar_implementation_D0, points: 10, grade: true, solution: false -->
```python

```

## Task 2.5

**[5 Point(s)]**

## Evaluation A) [5 Points]

Now we want to compare the performance of the different solvers. For that, run every combination of `algorithm` x `heuristic` with the puzzle configurations
1. `[0, 1, 2, 3, 4, 5, 8, 6, 7]`
1. `[8, 7, 6, 5, 1, 4, 2, 0, 3]`
1. `[1, 5, 7, 3, 6, 2, 0, 4, 8]`

for each run, record the `expanded nodes`, `execution time` and the `path cost`.
```
[
  {
    'puzzle': 1,
    'algorithm': 'astar',
    'heuristic': 'manhattan',
    'expanded nodes': 0,
    'execution time': 0.0,
    'path cost': 0
  }
]
```

<!-- grade_id: Astar_implementation_E, points: 0, grade: false, solution: true -->
```python
evaluation_results = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_Astar_implementation_E0, points: 5, grade: true, solution: false -->
```python

```

## Task 2.6

**[5 Point(s)]**

## Evaluation B) [5 Points]
Lastly, we want to get an estimate on how well the solvers are performing in general. For that, we first have to define a measue of *difficulty* of an initial configuration and whether ther exists a solution at all. 

### Inversions and Solvability of the 8 and 15 Puzzle
The difficulty of an 8 puzzle configuration can be acessed by counting the number of inversions it contains with respect to the goal configuration. If the square containing the number $i$ appears "before" (reading the squares in the box from left to right and top to bottom) $n$ numbers that are less than $i$, then call it an inversion of order $n$, and denote it $n_i$. Then define
$$N=\sum_{i=1}^8n_i = \sum_{i=2}^8n_i$$
where the sum only needs to run from 2 to 8, since there are no numbers less than 1 (so $n_1$ must equal 0). Stated more simply, $N = i(p)$ is the number of *permuted inversions* in the list of numbers.

For the **15 puzzle** there actually exists a proof, which in its principles also holds for the 8 puzzle in modified form:  
We additionaly define $e$ to be the row number of the empty square. Then if $N + e$ is even, then the configuration is solvable, otherwise it is not. In other words, if the permutation symbol $(-1)^{i(p)}$ of the list is $+1$, the configuration is solvable, whereas if the signature is -1, it is not. This can be formally proved using *alternating groups*.

For more details, check out this article by [MathWorld](https://mathworld.wolfram.com/15Puzzle.html)

> While odd permutations of the puzzle are impossible to solve (Johnson 1879), all even permutations are solvable (Story 1879).

With those definitions it now becomes possible to evaluate the general performance of the solvers with respect to the complexity of the initial configuration.

1. Compelte the `inversion` function below, such that it calculates the total number of inversions in a puzzle configuration.
1. Use the `inversion` function in order to complete the `is_solvable` function, which should return either `True` or `False` depending on whether a configuration is solvable or not.
1. Generate 100 random **solvable** puzzle configurations
1. Solve each configuration with every combination of `algorithm` and `heuristic` and again record the `expanded nodes`, `execution time` and the `path cost`, as well as the `complexity`, which is the total inversion count of the initial configuration.
1. Plot the metrics over the `complexity` of the puzzle's initial configuration

<!-- grade_id: Astar_implementation_F, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot

def inversion(configuration: List[int]) -> int:
    """
    Finds all inversions in a puzzle configuration and returns their total number
    :param configuration: the configuration of the 8 puzzle to count its inversions
    :returns: number of inversions in the puzzle configuration
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def is_solvable(configuration: List[int]) -> bool:
    """
    Checks whether a given puzzle configuration is solvable or not
    :param configuration: the initial configuration of the puzzle to check for solvability
    :returns: True if the configuration is solvable; False otherwise
    """
    # YOUR CODE HERE
    raise NotImplementedError()


# YOUR CODE HERE
raise NotImplementedError()
```

### Test

```python
from typing import List

class Puzzle:

    def __init__(self, init_state: List[int], puzzle_type: int = 8):
        """
        Initialize a new 8-puzzle board
        :param init_state: Initial configuration of the board
        """
        self.init_state = init_state
        self.goal_state = [i for i in range(1, 9)] + [0]
        self.explored_states = set()

    def goal_test(self, state: List[int]) -> bool:
        """Test if goal state is reached
        :param state: board configuration to check
        :return: True if the passed configuration is equal to goal configuration, False otherwise
        """
        return state == self.goal_state

    def is_explored(self, state: List[int]) -> bool:
        """Check if a particular board configuration has already been explored
        :param state: board configuration to check
        :return: True if a particular configuration has already been explored, False otherwise
        """
        return tuple(state) in self.explored_states

    def reset(self):
        """Reset the puzzle to its initial state"""
        self.explored_states.clear()
```

<!-- grade_id: test_Astar_implementation_F0, points: 5, grade: true, solution: false -->
```python

def move_left(position: int) -> int:
    """Move one position left in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving to the left
    """
    return position - 1 if position % 3 != 0 else -1

def move_right(position: int) -> int:
    """Move one position right in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving to the right
    """
    return position + 1 if position % 3 != 2 else -1

def move_up(position: int) -> int:
    """Move one position up in 8 puzzle if possible
    :param position: current position of the 0 tile
    :return: new position of the 0 tile after moving upwards
    """
    return position - 3 if position >= 3 else -1

def move_down(position: int) -> int:
    """Move one position down in 8 puzzle if possible
    :param position: current position of the 0 tile.
    :return: new position of the 0 tile after moving downwards
    """
    return position + 3 if position < 6 else -1

def get_possible_moves(state: List[int]) -> List[List[int]]:
    """Check whether a move is possible in left, right, up, down direction and store it.
    :param state: current configuration of the puzzle as one dimensional list
    :return: list containing the new configurations after applying all possible moves
    """
    possible_moves = []
    position = state.index(0)  # Find the position of the empty tile (0)
    if (new_position := move_left(position)) != -1:
        new_state = state[:]
        new_state[position], new_state[new_position] = new_state[new_position], new_state[position]
        possible_moves.append(new_state)
    if (new_position := move_right(position)) != -1:
        new_state = state[:]
        new_state[position], new_state[new_position] = new_state[new_position], new_state[position]
        possible_moves.append(new_state)
    if (new_position := move_up(position)) != -1:
        new_state = state[:]
        new_state[position], new_state[new_position] = new_state[new_position], new_state[position]
        possible_moves.append(new_state)
    if (new_position := move_down(position)) != -1:
        new_state = state[:]
        new_state[position], new_state[new_position] = new_state[new_position], new_state[position]
        possible_moves.append(new_state)
    return possible_moves
```

```python
x = get_possible_moves([1,2,0,3,4,5,6,7,8])
print(x)
```

Output:
```
[[1, 0, 2, 3, 4, 5, 6, 7, 8], [1, 2, 5, 3, 4, 0, 6, 7, 8]]
```

```python
from typing import List

def hamming_distance(state: List[int]) -> int:
    """Calculate the Hamming distance for a particular configuration
    :param state: current configuration of the puzzle
    :return: the number of misplaced tiles in the given configuration
    """
    goal_state = [1, 2, 3, 4, 5, 6, 7, 8, 0]
    distance = 0
    for i in range(len(state)):
        if state[i] != goal_state[i]:
            distance += 1
    return distance

def manhattan_distance(state: List[int]) -> int:
    """Function to calculate the Manhattan distance for a
    particular configuration
    :param state: current configuration of the puzzle
    :return: the accumulated Manhattan distance between each tile and its goal position in the given configuration
    """
    goal_positions = {1: (0, 0), 2: (0, 1), 3: (0, 2),
                      4: (1, 0), 5: (1, 1), 6: (1, 2),
                      7: (2, 0), 8: (2, 1), 0: (2, 2)}  # Goal positions of each tile
    distance = 0
    for i in range(len(state)):
        current_row, current_col = divmod(i, 3)  # Calculate row and column of current tile
        goal_row, goal_col = goal_positions[state[i]]  # Get goal position of current tile
        distance += abs(current_row - goal_row) + abs(current_col - goal_col)  # Manhattan distance calculation
    return distance


```

```python
assert hamming_distance([0,1,2,3,4,5,8,6,7]) == 9
assert manhattan_distance([1,2,3,4,0,5,6,8,7]) == 8
```

```python
from typing import Callable, List, Tuple
from heapq import heappop, heappush

class Node:
    def __init__(self, state: List[int], g: int, f: int, parent=None):
        self.state = state
        self.g = g  # Cost to reach this state
        self.f = f  # Estimated total cost (g + h)
        self.parent = parent  # Parent node

    def __lt__(self, other):
        return self.f < other.f


def astar_search(board: Puzzle, heuristic: Callable) -> Tuple[List[List[int]], int]:
    """
    :param board: the 8-puzzle to solve
    :param heuristic: the heuristic function to use
    :return: an ordered list with the solution path and the number of total nodes expanded
    """
    visited_states = set()
    start_node = Node(board.init_state, 0, heuristic(board.init_state))
    open_list = [(start_node.f, start_node)]

    while open_list:
        _, current_node = heappop(open_list)
        current_state = current_node.state

        if board.goal_test(current_state):
            # Reconstruct the solution path
            solution_path = [current_state]
            while current_node.parent:
                current_node = current_node.parent
                solution_path.append(current_node.state)
            solution_path.reverse()
            return solution_path, len(visited_states)

        visited_states.add(tuple(current_state))

        for new_state in get_possible_moves(current_state):
            if tuple(new_state) not in visited_states:
                new_g = current_node.g + 1
                new_f = new_g + heuristic(new_state)
                new_node = Node(new_state, new_g, new_f)
                new_node.parent = current_node
                heappush(open_list, (new_node.f, new_node))

    return [], len(visited_states)
```

```python
p = Puzzle([0,1,2,3,4,5,8,6,7])
path, expanded_nodes = astar_search(p, manhattan_distance)
```

```python
print(len(path))
print(expanded_nodes)
```

Output:
```
23
1256
```